<font size="4" style="color:red;"> **IMPORTANT: ** When submitting this homework notebook, please modify only the cells that start with:</font>

```python
# modify this cell
```

**Note:**  notice that no packages are imported for this assignment. This is because you do not need any python packages.

# Probability Inequalities


For the binomial distribution $X\sim B_{p,n}$ with mean $\mu=np$ and variance $\sigma^2=np(1-p)$, 
we would like to upper bound the probability $P(X\ge c\cdot \mu)$ for $c\ge1$. 
The lectures introduced three bounds:

Markov: $$P(X\ge \alpha\mu)\le \frac{1}{\alpha},\quad\quad\forall \alpha\ge 1,$$
Chebyshev: $$P(|X-\mu|\ge \alpha\sigma)\le \frac{1}{\alpha^2},\quad\quad \forall \alpha\ge 1,$$
Note that, while double-sided, this inequality also bounds $P(X\ge\mu+\alpha)$
$$P(X\ge \mu+\alpha\sigma)\le P(|X-\mu|\ge \alpha\sigma)\le \frac{1}{\alpha^2},$$
Chernoff: $$P(X\ge (1+\delta)\mu)\le e^{-\frac{\delta^2}{2+\delta}\mu},\quad\quad\forall \delta\ge0.$$


```python
from math import exp
from math import sqrt
```

import exponential function exp from math

## Problem 1

Write three functions **Markov**, **Chebyshev** and **Chernoff** that take $n$, $p$ and $c$ as inputs and return the upper bounds for $P(X\ge c\cdot np)$ given by the above Markov, Chebyshev, and Chernoff inequalities as outputs.

<font  style="color:blue"> **Code:**</font>
```python
print Markov(100.,0.2,1.5)
print Chebyshev(100.,0.2,1.5)
print Chernoff(100.,0.2,1.5)
```


<font  style="color:magenta"> **Output**</font>
```
0.6666666666666666
0.16
0.1353352832366127

```

Markov: $$P(X\ge \alpha\mu)\le \frac{1}{\alpha},\quad\quad\forall \alpha\ge 1,$$<br>
$$P(X\ge a)\le \frac{\mu}{a},\quad\quad where \quad a =\alpha\mu \quad\forall \alpha\ge 1,$$

```python

# modify this cell

def Markov(n, p, c):
    # inputs: 3 floats as described above
    # output: a variable of type float
    
    #
    # YOUR CODE HERE
    #
    return 1/c
```

Chebyshev: $$P(|X-\mu|\ge \alpha\sigma)\le \frac{1}{\alpha^2},\quad\quad \forall \alpha\ge 1,$$
Note that, while double-sided, this inequality also bounds $P(X\ge\mu+\alpha)$
$$P(X\ge \mu+\alpha\sigma)\le P(|X-\mu|\ge \alpha\sigma)\le \frac{1}{\alpha^2},$$
$$P(X\ge c\mu) = P(X\ge \mu+\alpha\sigma)\le P(|X-\mu|\ge \alpha\sigma)\le \frac{1}{\alpha^2},$$

```python

# modify this cell

def Chebyshev(n, p, c):
    # inputs: 3 floats as described above
    # output: a variable of type float
    
    #
    # YOUR CODE HERE
    #
    # c*mu = mu + alpha * std
    alpha = (c-1)*n*p / sqrt(n*p*(1-p))
    return 1/(alpha**2)
```

Chernoff: $$P(X\ge (1+\delta)\mu)\le e^{-\frac{\delta^2}{2+\delta}\mu},\quad\quad\forall \delta\ge0.$$

```python

# modify this cell

def Chernoff(n, p, c):
    # inputs: 3 floats as described above
    # output: a variable of type float
    
    #
    # YOUR CODE HERE
    #
    # c = 1 + delta
    std = c - 1
    return (exp(-(std**2)/(2+std)*n*p))
```

<!-- grade_id: ex1, points: 5, grade: true, solution: false -->
```python
assert abs(Markov(200.,0.25,1.25)-0.8)< 10**-5
assert abs(Chebyshev(100.,0.25,1.25)-0.48)< 10**-5
assert abs(Chernoff(100.,0.25,1.25)-0.4993517885992762)< 10**-5
#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

```python

```