<font size="4" style="color:red;"> **IMPORTANT: ** When submitting this homework notebook, please modify only the cells that start with:</font>

```python
# modify this cell
```

```python
from scipy import stats
from math import sqrt
```

# Confidence Interval

For a sample with size large enough, by central limit theorem we can assume its mean follows normal distribution. And if we also know the standard deviation of the population, we are able to calculate a confidence interval to estimate the population mean.

## Problem 1

A confidence interval is usually given by sample mean $m$ plus and minus a margin of error $r$:
$$[m-r,m+r]$$
The confidence interval is larger (less precise) for large confidence level, and smaller (more precise) for small confidence level.

For this problem, you are asked to write a function **Error** that calculates the margin of error $r$ given sample size $n$, confidence level $p$ and standard deviation of the population $s$.

<font  style="color:blue"> **Code:**</font>
```python
Error(40,0.95,20)

Error(40,0.95,10)       
```


<font  style="color:magenta"> **Output**</font>
```
6.1979503230456148
3.0989751615228074
```

```python
# modify this cell

def Error(n,p,s):
    # inputs: sample size n, confidence level p and standard deviation s
    # output: margin of error r
    
    #
    # YOUR CODE HERE
    #
    z_p = stats.norm.ppf((1 + p)/2)
    return z_p * s / sqrt(n)
```

<!-- grade_id: ex1, points: 5, grade: true, solution: false -->
```python
# Check Function
assert abs(Error(60,0.9,20)-4.2469938027546128) < 10**-5 
assert abs(Error(60,0.9,10)-2.1234969013773064) < 10**-5

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

# Problem 2

For this problem, you are asked to write a function **Confidence** that calculates the confidence level $p$ given sample size $n$, margin of error $r$ and standard deviation of the population $s$.

<font  style="color:blue"> **Code:**</font>
```python
Confidence(40,6,20)

Confidence(40,8,20)      
```


<font  style="color:magenta"> **Output**</font>
```
0.94222042887640267
0.98858796361399826
```

```python
# modify this cell

def Confidence(n,r,s):
    # inputs: sample size n, margin of error r, and standard deviation s
    # output: confidnce level r
    
    #
    # YOUR CODE HERE
    #
    z_p = r * sqrt(n) / s
    phi = stats.norm.cdf(z_p)
    # P(-a < Z < a) = 2*phi - 1
    return (2 * phi - 1)
```

<!-- grade_id: ex2, points: 5, grade: true, solution: false -->
```python
# Check Function
assert abs(Confidence(60,5,20)-0.94719248858388649) < 10**-5 
assert abs(Confidence(60,6,20)-0.97986324844965367) < 10**-5

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```