<font size="4" style="color:red;"> **IMPORTANT: ** When submitting this homework notebook, please modify only the cells that start with:</font>

```python
# modify this cell
```


# Different Dice


So far we mostly considered standard 6-faced dice. The following problems explore dice with different number of faces. 

```python
import numpy as np
%pylab inline
```

Output:
```
Populating the interactive namespace from numpy and matplotlib
```

## Problem 1

Suppose that a 6-sided die is rolled $n$ times. Let $X_i$ be the value of the top face at the $i$th roll,
and let $X\triangleq\max_{1\le i\le n} X_i$ be the highest value observed. For example, if $n=3$ and the three
rolls are 4, 1, and 4, then $X_1=4, X_2=1, X_3=4$ and $X=4$. 

To find the distribution of $X$, observe first that $X\le x$ iff $X_i\le x$ for all $1\le i\le n$,
hence $P(X\le x)=(x/6)^n$. It follows that $P(X=x)=P(X\le x)-P(X\le x-1)=(x/6)^n-((x-1)/6)^n$.
For example, $P(X=1)=(1/6)^n$, and $P(X=2)=(1/3)^n-(1/6)^n$.

In this problem we assume that each of the $n$ dice has a potentially different number of faces, denoted $f_i$,
and ask you to write a function **largest_face** that determines the probability $P(x)$ that the highest top face observed is $x$. **largest_face** takes a vector $\boldsymbol f$ of positive integers, interpreted as the number of faces of each of the dice, and a value $x$ and returns $P(x)$. For example, if $\boldsymbol f=[2, 5, 7]$, then three dice are rolled, and $P(1)=(1/2)\cdot(1/5)\cdot(1/7)$ as all dice must turn up 1, while $P(7)=1/7$ as the third die must turn up 7. In addition, $P(2) = (2/2)\cdot(2/5)\cdot(2/7) - (1/2)\cdot(1/5)\cdot(1/7)$, which is the probability that all the three dices should be at most 2, while at least one of them is 2 (we minus the all 1 case).

In all, the function `largest_face(f, x_max)` asks you to roll a list of dice provided by `f`, and to calculate the probability that the largest number among the outcomes is `x_max`.

<font  style="color:blue">* **Sample run** *</font>
```python
print(largest_face([2,5,8],8))
print(largest_face([2], 1))
largest_face( [3,4], 2)
print(largest_face([2, 5, 7, 3], 3))
```


<font  style="color:magenta">* **Expected Output** *</font>
```
0.125
0.5
0.25
0.180952380952
```

```python
# modify this cell

def largest_face(f, x_max):
    # inputs: m is a list of integers and m_max is an integer
    # output: a variable of type 'float'
    
    #
    # YOUR CODE HERE
    #
    p1 = 1
    p2 = 1

    for f_i in f:
        if x_max > f_i:
            continue
        else: # P(X=x) = P(X<=x) - P(X<=x-1)
            p1 = p1 * x_max/f_i
            p2 = p2 * (x_max - 1)/f_i
    return p1 - p2
```

<!-- grade_id: ex_1, points: 5, grade: true, solution: false -->
```python
# Check Function

assert abs( largest_face([5],3) - 0.19999999999999996 ) < 10**-5
assert abs( largest_face( [11,5,4], 5) - 0.16363636363636364  ) < 10**-5
assert abs( largest_face(range(1,10), 3) - 0.011348104056437391 ) < 10**-5 


#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

## Problem 2

Write a function **face_sum** that takes a vector $\boldsymbol f$ that as in the previous problem represents the number of faces of each die,  and a positive integer $s$, and returns the probability that the sum of the top faces observed is $s$. For example, if $\boldsymbol f=[3, 4, 5]$ and $s\le 2$ or $s\ge 13$, **face_sum** returns 0, and if $s=3$ or $s=12$, it returns $(1/3)\cdot(1/4)\cdot(1/5)=1/60$.

Hint: The **constrained-composition** function you wrote for an earlier probelm may prove handy. 

<font  style="color:blue"> * **Sample run** *</font>
```python
print(face_sum([3, 4, 5], 13))
print(face_sum([2,2],3))
print(face_sum([3, 4, 5], 7))
```


<font  style="color:magenta"> * **Expected Output** *</font>
```
0.0
0.5
0.18333333
```

### Helper Code

Below is a correct implementation of **constrained_composition**. Call this function in your implementation of **face_sum**.

##### From HW 4

Let $\boldsymbol{m}=(m_1,...,m_k)$ be a vector of $k$ positive integers. A $k$-composition $(a_1,...,a_k)$ of $n$ is $\boldsymbol{m}$-constrained if $a_i\leq m_i$ for all $1\le i\le k$. For example, $(1,1,3)$ and $(2,1,2)$ are the only $(2,1,4)$-constrined 3-partitions of 5. 

A function **constrained_compositions** that takes a natural number $n$ and a vector $\boldsymbol{m}$ of $k$ positive integers, and prints the set of all $\boldsymbol m$-constrained $k$-compositions of $n$. Note that $k$ can be inferred from $\boldsymbol m$. 

<font  style="color:blue"> **Code**</font>
```python
func_out = constrained_compositions(7, [3, 2, 5])

print "all possible combinations: "
print_compositions(func_out)
print;print "Actual Output from compositions(3,7)"; func_out
```


<font  style="color:magenta"> **Output**</font>
```
all possible combinations: 
1 + 2 + 4 
2 + 1 + 4 
3 + 2 + 2 
3 + 1 + 3 
1 + 1 + 5 
2 + 2 + 3 

Actual Output from compositions(3,7)
{(1, 1, 5), (1, 2, 4), (2, 1, 4), (2, 2, 3), (3, 1, 3), (3, 2, 2)}
```

```python
def constrained_compositions(n, m):
    # inputs: n is of type 'int' and m is a list of integers
    # output: a set of tuples
    
    k = len(m)
    parts = set()
    if k == n:
        if 1 <= min(m):
            parts.add((1,)*n)
    if k == 1:
        if n <= m[0]:
            parts.add((n,))
    else:
        for x in range(1, min(n-k+2,m[0]+1)):
            for y in constrained_compositions(n-x, m[1:]):
                parts.add((x,)+y)
    return parts
```

### exercise:

```python
# modify this cell

def face_sum(m, s):
    # inputs: m is list of integers and s is an integer
    # output: a variable of type 'float'
    
    #
    # YOUR CODE HERE
    #
    p = 1
    if s > sum(m) or s < len(m):
        return 0
    else:
        for m_i in m:
            p = p * 1/m_i
        return p*len(constrained_compositions(s,m))      
        
```

<!-- grade_id: ex_2, points: 5, grade: true, solution: false -->
```python
# Check Function
assert abs( face_sum([2, 2],2) - 0.25  ) < 10**-5
assert abs( face_sum([2, 2],10) - 0.0  ) < 10**-5
assert abs( face_sum(range(1, 10),20) - 0.03037092151675485  ) < 10**-5

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

```python







```